# Working with a database

This week's assignment has a few basic steps.  First, we're going to pull some data down off the internet and store it into our MySQL database.  Make sure that you use your username as part of the table name as show in the examples so that you don't create a problem for other students.

Then, we'll merge that with some data already in the database and calculate a few results.  When it comes to calculating the results, you can do so either with SQL or with Pandas operations.


## PART 1: Setup your database connection and table name

In the code below, change the value of the variable `MYTABLE` to use your own username rather that `'pboal'`

You can then use `MYTABLE` in the rest of your code to reference that table name.

```python
import numpy as np
import os
import pymysql
from sqlalchemy import create_engine
import pandas as pd
import getpass
```

```python
# connection details of the server
database = 'hds5210'
port = '3306'
host = 'slucor2022-instance-1.cgdcoitnku0k.us-east-1.rds.amazonaws.com'
user_name = 'slucor2022'
password = 'SLUcor2022'

```

```python
conn = create_engine('mysql+pymysql://' + user_name + ':' + password + '@' + host + '/' + database)
```

```python
## create mytable

myname = getpass.getuser().split('-')[1]
myname
MYTABLE= myname + '_data'
```

<!-- grade_id: q01-tests, points: 5, grade: true, solution: false -->
```python
assert(MYTABLE != 'pboal')
assert(conn.connect())
```

## PART 2: Bring in outside data

Grab the data from this URL and put it into a database table named with your `username_data`.

https://opendata.arcgis.com/api/v3/datasets/ed0bc689bde246b18835c7529ba4efb4_0/downloads/data?format=csv&spatialRefId=4326

By the end of your cell, the table should be created.  The tests are going to verify that the table exists and looks right.

<!-- grade_id: q02-solution, points: 0, grade: false, solution: true -->
```python
## read the data

df = pd.read_csv('Electronic_Health_Record_EHR_Incentive_Program_Payments_to_Eligible_Hospitals.csv')
df.to_sql(myname + '_data', conn, index=True, if_exists='replace')
df
```

Output:
```
X          Y  OBJECTID  \
0   -124.142009  40.783559         1   
1   -122.086674  37.632915         2   
2   -122.295861  38.325402         3   
3   -120.766467  38.350920         4   
4   -122.435820  37.769049         5   
..          ...        ...       ...   
318 -122.173214  37.436328       319   
319 -118.186632  33.808293       320   
320 -121.455450  38.553012       321   
321 -118.286401  34.067178       322   
322 -117.152024  32.798322       323   

                                         Provider_Name         NPI    CCN  \
0             ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1609858950  50006   
1                             HAYWARD SISTERS HOSPITAL  1942298153  50002   
2             ST JOSEPH HEALTH NORTHERN CALIFORNIA LLC  1235218785  50009   
3                              SUTTER VALLEY HOSPITALS  1447494323  50014   
4                                 Sutter Bay Hospitals  1659439834  50008   
..                                                 ...         ...    .
```

```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
dxyz.shape
```

Output:
```
(323, 22)
```

<!-- grade_id: q02-tests, points: 5, grade: true, solution: false -->
```python
dxyz = pd.read_sql_query('SELECT * FROM ' + MYTABLE, conn)
assert(dxyz.shape == (323,22))
assert(list(dxyz.columns) == ['index', 'X', 'Y', 'OBJECTID', 'Provider_Name', 'NPI', 'CCN',
       'Business_Street_Address', 'Business_City', 'Business_County',
       'Business_ZIP_Code', 'Business_State_Territory', 'Payment_Year_Number',
       'Program_Type', 'Medicaid_EP_Hospital_Type', 'total_payments',
       'Last_Payment_Criteria', 'Recent_Disbursement_Amount', 'Latitude',
       'Longitude', 'Last_Program_Year', 'Last_Payment_Year'])
```

## PART 3: Combine with other data in the database

In the database, there is an existing table called `population`.  We want to merge the DHCS datafile loaded above with the population data available in this other database table  The tables can be merged on `MYTABLE`'s `Business_ZIP_Code` field and `population`'s `Zip` field.

Note that not all `Zip_Codes` from your downloaded file have to be in the `population` table.  If they aren't, then I want you to eliminate the non-matching records.  That is, only keep the records that have a matching ZIP code in both sets of data.

Answer the question:
Which providers are located in the zipcode with the largest population?

Put your answer in the form `answer = ['a', 'list', 'of', 'NPI', 'like', '1593042103]`

```python
table_name = {}
DF = pd.read_sql_query('SELECT * FROM ' + "population", conn)
for num,table_ in pd.read_sql_query("SHOW TABLES", conn).iterrows():
    name = str(table_['Tables_in_hds5210'])
    table_name[name] = pd.read_sql("SELECT * FROM "+name,conn)
population = table_name['population']
```

```python
population
```

Output:
```
index    Zip  Population
0         0  90011      110750
1         1  90650      104524
2         2  91331      103683
3         3  90201      100512
4         4  90044       99980
...     ...    ...         ...
1645   1645  95699          13
1646   1646  96119          11
1647   1647  92304           7
1648   1648  93530           6
1649   1649  92364           5

[1650 rows x 3 columns]
```

```python
answer = pd.merge(table_name[MYTABLE],population,left_on = 'Business_ZIP_Code',right_on = 'Zip',how = 'inner')
maximum = answer[(answer['Population'] ==answer['Population'].max())]
```

```python
answer = [str([maximum[['NPI']]][0]).split(" ")[14]]
```

<!-- grade_id: q03-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == list)
assert(answer == ['1194016923'])

```

## PART 4: Total by hospital type

This is a multistep process:
* From our downloaded data file, compute the `Total payments` per ZIP code and Medicaid EP Hospital Type.
* Then merge that with the `population` data to compute a `Total payments` per person.
* Then average that across all of the `Medicaid EP Hospital Types` to get an average per persona payment for these type of hospital.

Your answer should be in structure of a data frame with at least two columns:
* Medicaid_EP_Hospital_Type
* Avg_Pay_per_Capita

## First using SQL

<!-- grade_id: q04-solution, points: 0, grade: false, solution: true -->
```python
total_payment = pd.read_sql_query("""
  SELECT  a.Medicaid_EP_Hospital_Type, sum(a.total_payments) as payments, a.Business_ZIP_Code
  
  FROM 
    shiv80_data a 
  GROUP BY 
     a.Medicaid_EP_Hospital_Type, a.Business_ZIP_Code
""", conn)
```

```python
total_payment.to_sql('payment_data', conn, index=False, if_exists='replace')
```

```python
avg_pay = pd.read_sql_query("""
  SELECT  a.Medicaid_EP_Hospital_Type, AVG(a.payments/p.Population) as Avg_Pay_per_Capita
  
  FROM 
    population p  JOIN
    payment_data a ON p.Zip = a.Business_ZIP_Code
  GROUP BY 
    a.Medicaid_EP_Hospital_Type
  
""", conn)
```

```python
avg_pay
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals          106.582669
1      Children's Hospitals          329.550946
```

<!-- grade_id: q04-tests, points: 5, grade: true, solution: false -->
```python
assert(type(answer) == pd.core.frame.DataFrame)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Acute Care Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 106.583)
assert(round(answer.query("Medicaid_EP_Hospital_Type == 'Children\\'s Hospitals'")['Avg_Pay_per_Capita'].sum(),3) == 329.551)

```

```python
answer
```

Output:
```
Medicaid_EP_Hospital_Type  Avg_Pay_per_Capita
0      Acute Care Hospitals          106.582669
1      Children's Hospitals          329.550946
```

## Now using Pandas instead of SQL

```python
pdf = df.pivot_table(values = 'total_payments',index= ['Business_ZIP_Code', 'Medicaid_EP_Hospital_Type'],aggfunc = np.ma.sum,fill_value = 0)
pdf.reset_index()
```

Output:
```
Business_ZIP_Code Medicaid_EP_Hospital_Type  total_payments
0                90004      Acute Care Hospitals      1093021.73
1                90012      Acute Care Hospitals      3614739.93
2                90015      Acute Care Hospitals     10045484.04
3                90017      Acute Care Hospitals      3899845.21
4                90020      Children's Hospitals       427640.22
..                 ...                       ...             ...
262              96104      Acute Care Hospitals       184615.39
263              96122      Acute Care Hospitals       531907.08
264              96130      Acute Care Hospitals      1371962.88
265              96150      Acute Care Hospitals       786769.13
266              96161      Acute Care Hospitals       621315.39

[267 rows x 3 columns]
```

```python
pdf2 = pd.merge(pdf.reset_index(),population,left_on = 'Business_ZIP_Code',right_on = 'Zip',how = 'inner')
pdf2['total_per_person'] = (pdf2["total_payments"])/pdf2['Population']
```

```python
answer = pdf2.pivot_table(values = 'total_per_person',index= 'Medicaid_EP_Hospital_Type',aggfunc = np.ma.mean,fill_value = 0)
answer.rename(columns={'total_per_person':'Avg_Pay_per_Capita'},inplace=True)
```



---



## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git pull
    !git add week15_assignment.ipynb
    !git commit -a -m "Submitting the week 15 programming assignment"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```

```python

```